---
title       : Insert the chapter title here
description : Insert the chapter description here
attachments :

--- type:RStudioMultipleChoiceExercise xp:50 skills:1 key:02d4222b54
## Lab Question 1

Create a new data frame that includes flights headed to SFO in February, and save 
this data frame as `sfo_feb_flights`. How many flights meet these criteria? 

*** =instructions
- 68
- 1345
- 2286
- 3563
- 32735

*** =hint
Use similar code as the most recent code block above

*** =pre_exercise_code
```{r,eval=FALSE}
# created Lab2.Rmd
```

*** =sct
```{r,eval=FALSE}
msg1 <- "Awesome!"
msg2 <- "Try again, fill in the code block with the answer using dplyr's `filter`."

test_mc(3, feedback_msgs = c(msg2, msg2, msg1, msg2, msg2))
```

*** =attachments
Lab2.Rmd: https://s3.amazonaws.com/assets.datacamp.com/production/course_1156/datasets/Lab2.Rmd
