--- type:RStudioMultipleChoiceExercise xp:50 skills:1 key:8582d21328
## Test auto open

This is a test to auto open .Rmd and history.


*** =instructions
- "Test1"
- "Test2"

*** =hint
Use the arrows on the plots pane to toggle between multiple plots!

*** =pre_exercise_code
```{r,eval=FALSE}

```

*** =sct
```{r,eval=FALSE}
msg1 <- "Correct!"
msg2 <- "Incorrect. Look at the history pane and run the last command!"

test_mc(1, feedback_msgs = c(msg1, msg2))
```

*** =attachments
.Rprofile: https://s3.amazonaws.com/assets.datacamp.com/production/course_1161/datasets/lab1_ex7.Rprofile
.init.R: https://s3.amazonaws.com/assets.datacamp.com/production/course_1161/datasets/lab1_ex7.init.R
lab1_ex7.Rmd: https://s3.amazonaws.com/assets.datacamp.com/production/course_1161/datasets/lab1_ex7.Rmd

